# Assignment 1

* <b>Do not change the structure</b> of the notebook, only fill in the parts where your code needs to be filled in.

--- 

* For your notebook, <b>use as filename</b> `assignment01_vunetID.ipynb`, where `vunetID` is your three letter, three digit VU ID.

```python
# read neccessary libraries
import numpy as np
import pandas as pd
pd.options.display.max_rows = 20
np.set_printoptions(precision = 4, suppress = True)
```

### Manufacturing processes
The file below contains data about a manufacturing process. In this dataset, biological material and certain manufacturing parameters influence the yield. **Load** the CSV file and **calculate the average yield** across all records. Assign the result to `avg_yield`.  

<!-- grade_id: yield, points: 0, grade: false, solution: true -->
```python
file = 'manufacturing.csv'

### BEGIN SOLUTION
df = pd.read_csv(file)
avg_yield = df['Yield'].mean()
### END SOLUTION
```

<!-- grade_id: yield_correct, points: 1, grade: true, solution: false -->
```python
print(avg_yield)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (avg_yield, 40.17653409090909, 2)
### END HIDDEN TESTS
```

Output:
```
40.17653409090909
```

Some rows have one or more missing values. **Count how many such incomplete records (rows)** there are. Assign this total to `num_missing`.  

<!-- grade_id: missing, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
num_missing = df.shape[0] - df.dropna().shape[0]
### END SOLUTION
```

<!-- grade_id: missing-correct, points: 1, grade: true, solution: false -->
```python
print(num_missing)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (num_missing, 24, 2)
### END HIDDEN TESTS
```

Output:
```
24
```

The column **ManufacturingProcess03** contains some missing values. **replace any missing values with `1.0`**, then **calculate the mean** of the updated column. Store the result in `avg_mp3`.

<!-- grade_id: replace, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
newdata = df['ManufacturingProcess03']
newdata = newdata.fillna(1)
avg_mp3 = newdata.mean()
### END SOLUTION
```

<!-- grade_id: replace-correct, points: 1, grade: true, solution: false -->
```python
print(avg_mp3)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (avg_mp3, 1.4935795454545455, 2)
### END HIDDEN TESTS
```

Output:
```
1.4935795454545455
```

### Population of Amsterdam in 2024
The URL listed below contains tables on the demographics of Amsterdam. **Use `pandas.read_html()`** (or other appropriate methods) to read the table(s) from the webpage where you have to **locate the population estimate for Amsterdam on 2024-01-01** and assign that value to `population`.  

<!-- grade_id: html-q1, points: 0, grade: false, solution: true -->
```python
url = 'https://www.citypopulation.de/en/netherlands/admin/noord_holland/0363__amsterdam/'

### BEGIN SOLUTION
tables = pd.read_html(url)
statistics = tables[0]
population = statistics.iloc[0]['Population Estimate 2024-01-01']
### END SOLUTION
```

<!-- grade_id: html-q1-correct, points: 1, grade: true, solution: false -->
```python
print(population)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (population, 931298, 2)
### END HIDDEN TESTS
```

Output:
```
931298
```

 **Determine the fraction** that Amsterdam’s population (for 2024-01-01) represents of the **total population** listed in the same table. Store this fraction as `fraction`.

<!-- grade_id: html-q2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
fraction = population / statistics.loc[1]['Population Estimate 2024-01-01']
### END SOLUTION
```

<!-- grade_id: html-q2-correct, points: 1, grade: true, solution: false -->
```python
print(fraction)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (fraction, 0.05190330548914442, 4)
### END HIDDEN TESTS
```

Output:
```
0.05190330548914442
```

### Food database
The file listed below contains a database on food products. **Read** the file into a pandas structure. **Find out how many unique groups** (`group` field) exist in this dataset. Assign that number to `num_categories`.

<!-- grade_id: unique-categories, points: 0, grade: false, solution: true -->
```python
import json
file = 'database.json'

### BEGIN SOLUTION
# result = pd.read_json(file)
# num_categories = result['group'].drop_duplicates().count()
df = pd.read_csv(file)
num_categories = df['group'].nunique()
### END SOLUTION
```

<!-- grade_id: unique-categories-correct, points: 1, grade: true, solution: false -->
```python
print(num_categories)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (num_categories, 25, 2)
### END HIDDEN TESTS
```

From the same data, **count how many products** belong to the `Baby Foods` group. Store that count in `baby_total`.

<!-- grade_id: food-database, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
result = pd.read_json(file)
baby_total = (result['group'] == "Baby Foods").sum()
### END SOLUTION
```

<!-- grade_id: food-database-correct, points: 2, grade: true, solution: false -->
```python
print(baby_total)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (baby_total, 209, 2)
### END HIDDEN TESTS
```

Output:
```
209
```

### Bus database
The following XML file contains records on the LI Bus service. Read in the XML and **count how many records** (rows) are in this data. Assign the result to `num_records`.

<!-- grade_id: bus, points: 0, grade: false, solution: true -->
```python
from lxml import objectify
file = 'libus.xml'

### BEGIN SOLUTION
parsed = objectify.parse(open(file))
root = parsed.getroot()

data = []

for elt in root.INDICATOR:
    el_data = {}
    for child in elt.getchildren():
        el_data[child.tag] = child.pyval
    data.append(el_data)
    
perf = pd.DataFrame(data)
num_records = perf.shape[0]
### END SOLUTION
```

<!-- grade_id: bus-correct, points: 2, grade: true, solution: false -->
```python
print(num_records)
### BEGIN HIDDEN TESTS
from numpy.testing import assert_almost_equal
assert_almost_equal (num_records, 480, 2)
### END HIDDEN TESTS
```

Output:
```
480
```